# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT) (http://www.econ.ku.dk/phdstudent/noack/)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Array{Int64,2}:
 3  1  3
 3  2  3
 2  1  4
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Array{Float64,1}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Array{Float64,1}:
 7.0
 8.0
 7.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 LinearAlgebra.Adjoint{Int64,Array{Int64,2}}:
 3  3  2
 1  2  1
 3  3  4
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 LinearAlgebra.Transpose{Int64,Array{Int64,2}}:
 3  3  2
 1  2  1
 3  3  4
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Array{Int64,2}:
 22  11  26
 11   6  13
 26  13  34
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
3-element Array{Float64,1}:
 0.9999999999999997
 1.0               
 1.0000000000000002
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Array{Float64,2}:
 0.6777     0.781782
 0.31349    0.49033 
 0.0699567  0.562507
```

```julia
Atall\b
```

Output:
```
2-element Array{Float64,1}:
 -4.431949390951858
 14.157647062824969
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Array{Float64,2}:
 0.360666  0.360666
 0.387361  0.387361
 0.904023  0.904023
```

```julia
rankdef\b
```

Output:
```
2-element Array{Float64,1}:
 5.44553531354096 
 5.445535313540961
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Array{Float64,2}:
 0.99309   0.807182  0.952277
 0.776093  0.918919  0.901754
```

```julia
Ashort\bshort
```

Output:
```
3-element Array{Float64,1}:
 -0.045209201461404  
  0.5011792662084933 
  0.25081440108688285
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Array{Int64,1}:
 1
 2
 3
```

```julia
using LinearAlgebra
```

```julia
dot_v = dot(v, v)
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
outer_v = v*v'
```

Output:
```
3×3 Array{Int64,2}:
 1  2  3
 2  4  6
 3  6  9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

```julia
cross_v = cross(v, v)
```

Output:
```
3-element Array{Int64,1}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

Please click on `Validate` on the top, once you are done with the exercises.